---
  title: "2"
  description: ""
---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 96283bb60e   
```

`@projector_key`
edfd4e9aba1e41917d1885afd72c577f
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: eae86458c3   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 03a4b8bc90   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: 29cbf59f17   
```





`@hint`


`@possible_answers`


`@feedbacks`


---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: ca77facb3e   
```





`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


***



```yaml
type: NormalExercise 
xp: 50 
key: 3d2598231d   
```





`@instructions`


`@hint`


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```



***



```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 8e634ec51a   
```





`@instructions`


`@hint`


`@sct`

```{python}

```



---

## Insert exercise title here

```yaml
type: BulletExercise 
xp: 100 
key: 46d88d78d0   
```





`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


***



```yaml
type: NormalExercise 
xp: 100 
key: 2b98d8536e   
```





`@instructions`


`@hint`


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


